<!-- grade_id: 1, points: 0, grade: false, solution: true -->
```python


def squares(n):
    """Compute the squares of numbers from 1 to n, such that the 
    ith element of the returned list equals i^2.
    
    """
    ### BEGIN SOLUTION
    if n < 1:
        raise ValueError("n must be greater than or equal to 1")
    return [i ** 2 for i in range(1, n + 1)]
    ### END SOLUTION


```

<!-- grade_id: 2, points: 0, grade: true, solution: false -->
```python
assert 1 == 1
```